Zebrafish live imaging reveals only around 2% rather than 50% of motor neurons die through apoptosis during early development

Author response:

The following is the authors’ response to the original reviews.

Reviewer 1:

We thank the reviewer for the time and effort in providing very useful comments and suggestions for our manuscript.

(1) The results do not support the conclusions. The main "selling point" as summarized in the title is that the apoptotic rate of zebrafish motorneurons during development is strikingly low (~2% ) as compared to the much higher estimate (~50%) by previous studies in other systems. The results used to support the conclusion are that only a small percentage (under 2%) of apoptotic cells were found over a large population at a variety of stages 24-120hpf. This is fundamentally flawed logic, as a short-time window measure of percentage cannot represent the percentage in the long term. For example, at any year under 1% of the human population dies, but over 100 years >99% of the starting group will have died. To find the real percentage of motorneurons that died, the motorneurons born at different times must be tracked over the long term or the new motorneuron birth rate must be estimated. A similar argument can be applied to the macrophage results. Here the authors probably want to discuss well-established mechanisms of apoptotic neuron clearance such as by glia and microglia cells.

We chose the time window of 24-120 hpf based on the following two reasons: 1) Previous studies showed that although the time windows of motor neuron death vary in chick (E5-E10), mouse (E11.5-E15.5), rat (E15-E18), and human (11-25 weeks of gestation), the common feature of these time windows is that they are all the developmental periods when motor neurons contact with muscle cells. The contact between zebrafish motor neurons and muscle cells occurs before 72 hpf, which is included in our observation time window of 24-120 hpf. 2) Zebrafish complete hatching during 48-72 hpf, and most organs form before 72 hpf. More importantly, zebrafish start swimming around 72 hpf, indicating that motor neurons are fully functional at 72 hpf. Thus, we are confident that this 24-120 hpf time window covers the time window during which motor neurons undergo programmed cell death during zebrafish early development. We have added this information to the revised manuscript.

We frequently used “early development” in this manuscript to describe our observation. However, we missed “early” in our title. We therefore have added this ket word of “early” in the title in the revised manuscript.

Previous studies in zebrafish have shown that the production of spinal cord motor neurons largely ceases before 48 hpf, and then the motor neurons remain largely constant until adulthood (doi: 10.1016/j.celrep.2015.09.050; 10.1016/j.devcel.2013.04.012; 10.1007/BF00304606; 10.3389/fcell.2021.640414). Our observation time window covers the major motor neuron production process. Therefore, we believe that neurogenesis will not affect our findings and conclusions.

We discussed the engulfment of dead motor neurons by other types of cells in the discussion section.

(2) The transgenic line is perhaps the most meaningful contribution to the field as the work stands. However, the mnx1 promoter is well known for its non-specific activation - while the images suggest the authors' line is good, motor neuron markers should be used to validate the line. This is especially important for assessing this population later as mnx1 may be turned off in mature neurons.

The mnx1 promoter has been widely used to label motor neurons in transgenic zebrafish. Previous studies have shown that most of the cells labeled in the mnx1 transgenic zebrafish are motor neurons. In this study, we observed that the neuronal cells in our sensor zebrafish formed green cell bodies inside of the spinal cord and extended to the muscle region, which is an important morphological feature of the motor neurons.

Reviewer 2:

We thank the reviewer for the time and effort in making very useful comments and suggestions for our manuscript.

The FRET-based programmed cell death biosensor described in this manuscript could be very useful. However, the authors have not considered what is already known about the development and programmed cell death of zebrafish spinal motor neurons, and potential differences between motor neuron populations innervating different types of muscles in different vertebrate models. Without this context, the application of their new biosensor tool does not provide new insights into zebrafish motor neuron programmed cell death. In addition, the authors have not carried out controls to show the efficacy and specificity of their morpholinos. Nor have they described how they counted dying motor neurons, or why they chose the specific developmental time points they addressed. These issues are addressed more specifically below.

(1) Lines 12-13: Previous studies in zebrafish showed death of identified spinal motor neurons.

Line 103: In Figure 2A the cell body in the middle is that of identified motor neuron VaP. VaP death has previously been described in several publications. The cell body on the right of the same panel appears to belong to an interneuron whose axon can be seen extending off to the left in one of the rostrocaudal axon bundles that traverse the spinal cord. Higher-resolution imaging would clarify this.

Lines 163-164: Is this the absolute number of motor neurons that died? How were the counts done? Were all the motor neurons in every segment counted? There are approximately 30 identifiable VaP motor neurons in each embryo and they have previously been reported to die between 24-36 hpf. So this analysis is likely capturing those cells.

Our study examined the overall motor neuron apoptosis rather than a specific type of motor neuron death, so we did not emphasize the death of VaP motor neurons. We agree that the dead motor neurons observed in our manuscript contain VaP motor neurons. However, there were also other types of dead motor neurons observed in our study. The reasons are as follows: 1) VaP primary motor neurons die before 36 hpf, but our study found motor neuron cells died after 36 hpf and even at 84 hpf (revised Figure 4A). 2) The position of the VaP motor neuron is together with that of the CaP motor neuron, that is, at the caudal region of the motor neuron cluster. Although it’s rare, we did observe the death of motor neurons in the rostral region of the motor neuron cluster (revised Figure 2C). 3) There is only one or zero VaP motor neuron in each motor neuron cluster. Although our data showed that usually one motor neuron died in each motor neuron cluster, we did observe that sometimes more than one motor neuron died in the motor neuron cluster (revised Figure 2C). We included this information in the revised discussion.

(2) Lines 82-83: It is published that mnx1 is expressed in at least one type of spinal interneuron derived from the same embryonic domain as motor neurons.

The mnx1 promoter has been widely used to label motor neurons in transgenic zebrafish. Previous studies have shown that most of the cells labeled in the mnx1 transgenic zebrafish are motor neurons. In this study, we observed that the neuronal cells in our sensor zebrafish formed green cell bodies inside of the spinal cord and extended to the muscle region, which is an important morphological feature of the motor neurons.

Furthermore, a few of those green cell bodies turned into blue apoptotic bodies inside the spinal cord and changed to blue axons in the muscle regions at the same time, which strongly suggests that those apoptotic neurons are not interneurons. Although the mnx1 promoter might have labeled some interneurons, this will not affect our major finding that only a small portion of motor neurons died during zebrafish early development.

(3) Lines 161-162: Although this may be the major time window of neurogenesis, there are many more motor neurons in adults than in larvae. Neither of these references describes the increase in motor neuron numbers over this particular time span, so the rationale for this choice is unclear.

Lines 168-171: It is known that later developing motor neurons are still being generated in the spinal cord at this time, suggesting that if there is a period of programmed cell death similar to that described in chick and mouse, it would likely occur later. In addition, most of the chick and mouse studies were performed on limb-innervating motor neurons, rather than the body wall muscle-innervating motor neurons examined here.

Lines 237-238: Especially since new motor neurons are still being generated at this time.

Previous studies have shown that the production of spinal cord motor neurons largely ceases before 48 hpf in zebrafish, and then the motor neurons remain largely constant until the adulthood (doi: 10.1016/j.celrep.2015.09.050; 10.1016/j.devcel.2013.04.012; 10.1007/BF00304606; 10.3389/fcell.2021.640414). Our observation time window covers the major motor neuron production process. Therefore, we believe that neurogenesis will not affect our data and conclusions.

The death of motor neurons in limb-innervating motor neurons has been extensively studied in chicks and rodents, as it is easy to undergo operations such as amputation. However, previous studies have shown this dramatic motor neuron death does not only occur in limb-innervating motor neurons but also occurs in other spinal cord motor neurons (doi: 10.1006/dbio.1999.9413). In our manuscript, we studied the naturally occurring motor neuron death in the whole spinal cord during the early stage of zebrafish development.

(4) Lines 184-187: Previous publications showed that death of VaP is independent of limitations in muscle innervation area, suggesting it is not coupled to muscle-derived neurotrophic factors.

Lines 328-334: There have been many publications describing appropriate morpholino controls. The authors need to describe their controls and show that they know that the genes they were targeting were downregulated.

For the morpholinos, we did not confirm the downregulation of the target genes. These morpholino-related data are a minor part of our manuscript and shall not affect our major findings. We have removed the neurotrophic factors and morpholino-related data in the revised manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: It is widely accepted that neurons will die through apoptosis if they cannot receive enough growth factors during development of vertebrates; however, there is still no real-time observation showing this dying process in live animals. Here, we generated sensor zebrafish achieving live imaging of motor neuron apoptosis at single-cell resolution. Using these sensor zebrafish, we observed for the first time that in an apoptotic motor neuron, caspase-3 activation occurred quickly within 5-6 min and at the same time between the cell body and axon. Interestingly, we found that only around 2% of motor neurons died during zebrafish development, which is much lower than the generally believed 50% cell death occurred in embryonic stage of vertebrates. Our data also showed that most of the apoptotic bodies of these dead motor neurons were not cleared by macrophages. These sensor zebrafish can serve as powerful tools to study motor neuron apoptosis in vivo.

Revised version: It is widely accepted that neurons will die through apoptosis if they cannot receive enough growth factors during development of vertebrates; however, there is still no real-time observation showing this dying process in live animals. Here, we generated sensor zebrafish achieving live imaging of motor neuron apoptosis at single-cell resolution. Using these sensor zebrafish, we observed for the first time that in an apoptotic motor neuron, caspase-3 activation occurred quickly within 5-6 min and at the same time between the cell body and axon. Interestingly, we found that only around 2% of motor neurons died during zebrafish early development, which is much lower than the generally believed 50% cell death occurred in embryonic stage of vertebrates. Our data also showed that most of the apoptotic bodies of these dead motor neurons were not cleared by macrophages. These sensor zebrafish can serve as powerful tools to study motor neuron apoptosis in vivo.

Introduction

Programmed cell death or apoptosis is a dynamic process that occurs in most tissues in live animals. In the nervous system, apoptosis plays an essential role in the regulation of neuronal cell numbers (Hollville, Romero, & Deshmukh, 2019). During embryonic development, a large number of neuronal cells die through apoptosis (Fuchs & Steller, 2011). In pathological states, such as Alzheimer’s disease, Parkinson’s disease, and motor neuron disease, the death of neuronal cells severely impairs the normal function of the nervous system (Butterfield & Halliwell, 2019; Michel, Hirsch, & Hunot, 2016; Moujalled, Strasser, & Liddell, 2021).

When apoptosis occurs, either through intrinsic or extrinsic pathways, it ultimately leads to the cascade activation of caspases and the degradation of cellular structures, which is no exception for neuronal cells (Blanquie et al., 2017; Heck et al., 2008; Hernandez-Baltazar, Mendoza-Garrido, & Martinez-Fong, 2013). Due to the unique and complex morphology of neuronal cells, especially long axons, many important issues related to neuronal cell apoptosis have not been fully addressed, including: (1) How to visualize neuronal cell death in live animals? (2) How fast caspase-3 can be activated within a single neuron during apoptosis? (3) Which part of a neuron (cell body vs. axon) will die first? To address these important questions, we need to develop novel tools that can monitor neuronal cell apoptosis in real time and in live animals.

Previously, we have developed a fluorescence resonance energy transfer (FRET)-based biosensor, named sensor C3, for detecting cancer cell apoptosis. Our in vitro results showed that sensor C3 could detect the apoptosis of cancer cells both in 2D and 3D cultures by changing its color from green to blue in real time (Anand, Fu, Teoh, & Luo, 2015; Fu, Peh, Ngan, Wei, & Luo, 2018; Hao, Huang, Wu, Peng, & Luo, 2023; K. Li, Wu, Zhou, Tong, & Luo, 2021; Luo, Vivian, Pu, & Chang, 2001; Tian, Ip, Luo, Chang, & Luo, 2007; Wu, Li, Yuan, & Luo, 2021; Yang, Jia, Zhao, & Luo, 2022). Recently, we showed that sensor C3 could also be used to indicate zebrafish skin cell apoptosis (Jia, Song, Huang, Ge, & Luo, 2020).

Earlier version: In the present study, using Tol2 transposon-based transgenic technology, we generated novel sensor zebrafish and achieved the real-time tracking of motor neuron apoptosis at single-cell resolution in live zebrafish. More importantly, using these sensor zebrafish, we were able to obtain novel insights into the spatiotemporal properties and occurring rates of motor neuron death during development.

Revised version: In the present study, using Tol2 transposon-based transgenic technology, we generated novel sensor zebrafish and achieved the real-time tracking of motor neuron apoptosis at single-cell resolution in live zebrafish. More importantly, using these sensor zebrafish, we were able to obtain novel insights into the spatiotemporal properties and occurring rates of motor neuron death during early development.

Results

Generation of transgenic zebrafish expressing sensor C3 in motor neurons

Sensor C3 is a fusion protein of cyan fluorescent protein (CFP) and yellow fluorescent protein (YFP), between which is an amino acid linker containing the caspase-3 cleavage site DEVD (Figure 1A). In live cells, when sensor C3 is excited with a 458 nm laser, due to the energy transfer between CFP and YFP, the emission light of CFP can excite the YFP to emit green fluorescence, which can be detected at the wavelength of 535 nm. In apoptotic cells, activated caspase-3 proteins cleave sensor C3 at the DEVD site; due to the disruption of the FRET effect, sensor C3 emits blue fluorescence, which can be detected at the wavelength of 480 nm (Figure 1B).

The live imaging of motor neurons in sensor zebrafish. (A) The structure of sensor C3 predicted by using Rosetta software. (B) The principle of sensor C3 for apoptosis detection. (C-D) The expression of sensor C3 in the motor neurons of Tg(mnx1:sensor C3) zebrafish. The bright field image and fluorescent image were merged in (C) to show the location of motor neurons and their axons. The anatomical structures and developmental time are indicated.

In our previous publications, we have conducted many experiments to validate the specificity and sensitivity of sensor C3 for detecting caspase-3 activation and apoptosis. First, we confirmed that there was a strong energy transfer from CFP to YFP in purified sensor C3 proteins when CFP was excited. Second, we showed that the cleavage to sensor C3 is specific to caspase-3 and requires the presence of the cleavage sequence of DEVD in sensor C3 proteins. Third, sensor C3 is highly sensitive to caspase-3 cleavage, as it could detect caspase-3 activation at a nanomolar concentration and in 5 min when sensor C3-labelled HeLa cells underwent UV irradiation-induced apoptosis (Luo et al., 2001).

We have confirmed that the apoptosis detected by sensor C3 was consistent with classical apoptotic assays, including chromatin condensation staining, caspase-3 activity assay, DNA fragmentation assay, and morphological changes such as cell shrinkage (Luo et al., 2001; Tian et al., 2007). Furthermore, we showed that sensor C3 could detect the apoptosis of various cancer cells and zebrafish skin cells by color change no matter the apoptosis was induced by UV irradiation, chemical drugs, fluidic shear stress, immune cells-mediated killing, or during normal development (Hao et al., 2023; Jia et al., 2020; K. Li et al., 2021; Luo et al., 2001; Tian et al., 2007; Yang et al., 2022). All aforementioned results show that sensor C3 is a very good tool for detecting apoptosis in live cells and animals.

Earlier version: To detect the death of motor neurons in live animals, we cloned sensor C3 under a motor neuron-specific promoter (Arkhipova et al., 2012) and used this construct to generate a transgenic line of sensor zebrafish Tg(mnx1:sensor C3) in which only the motor neurons specifically expressed sensor C3 proteins. After obtaining the sensor zebrafish, we first examined the morphology of motor neurons in zebrafish. At 33 hours post fertilization (hpf), we observed many green motor neurons inside the spinal cord (Figure 1C-D). More detailed observation revealed that approximately 10 motor neurons bundled together with their cell bodies located within the spinal cord, and their axons converged into one bundle to extend into individual somite, which serves as a functional unit for the development and contraction of muscle cells (Figure 1C-D).

Revised version: To detect the death of motor neurons in live animals, we cloned sensor C3 under a motor neuron-specific promoter (Arkhipova et al., 2012; Flanagan-Steet, Fox, Meyer, & Sanes, 2005; Jao, Appel, & Wente, 2012; Rosenberg, Wolman, Franzini-Armstrong, & Granato, 2012; Stil & Drapeau, 2016) and used this construct to generate a transgenic line of sensor zebrafish Tg(mnx1:sensor C3) in which the motor neurons specifically expressed sensor C3 proteins. After obtaining the sensor zebrafish, we first examined the morphology of motor neurons in zebrafish. At 33 hours post fertilization (hpf), we observed many green motor neurons inside the spinal cord (Figure 1C-D). More detailed observation revealed that approximately 10 motor neurons bundled together with their cell bodies located within the spinal cord, and their axons converged into one bundle to extend into individual somite, which serves as a functional unit for the development and contraction of muscle cells (Figure 1C-D).

Earlier version: Both the cell body and axon of motor neurons can undergo apoptosis during zebrafish development

Revised version: Both the cell body and axon of motor neurons can undergo apoptosis during zebrafish early development

Earlier version: To determine the apoptotic status of motor neurons during the development, we took the images of CFP and YFP separately from the Tg(mnx1:sensor C3) zebrafish and merged them into FRET images. The green color will indicate the motor neurons are alive, and the blue color will indicate the neurons are apoptotic. The FRET imaging analysis showed that around 30 hpf during zebrafish early development, the cell bodies of some motor neurons appeared in blue. In contrast, most of the cell bodies of these motor neurons appeared in green. Furthermore, the blue signals were also observed in the axon bundle of the motor neurons (Figure 2A-C). Image stacks of the apoptotic motor neuron in Figure 2C at different optical section levels better illustrated the blue signals in the axon region (Figure 2D). This observation indicates that the majority of the motor neurons are alive, and the minority can undergo apoptosis. Furthermore, both the cell body and axon of the motor neurons can undergo apoptosis during the early development of zebrafish.

Revised version: To determine the apoptotic status of motor neurons during the development, we took the images of CFP and YFP separately from the Tg(mnx1:sensor C3) zebrafish and merged them into FRET images. The green color will indicate the motor neurons are alive, and the blue color will indicate the neurons are apoptotic. The FRET imaging analysis showed that around 30 hpf during zebrafish early development, the cell bodies of some motor neurons appeared in blue. In contrast, most of the cell bodies of these motor neurons appeared in green. Furthermore, the blue signals were also observed in the axon bundle of the motor neurons (Figure 2A-D). Figure 2C showed a motor neuron cluster inside of which more than one motor neuron died. Image stacks of the apoptotic motor neuron in Figure 2D at different optical section levels better illustrated the blue signals in the axon region (Figure 2E). This observation indicates that the majority of the motor neurons are alive, and the minority can undergo apoptosis. Furthermore, both the cell body and axon of the motor neurons can undergo apoptosis during the early development of zebrafish.

Earlier version: Both the cell body and axon of motor neurons can undergo apoptosis during zebrafish development. (A-C) FRET images of three somites showing motor neurons underwent apoptosis during the development of Tg(mnx1:sensor C3) zebrafish. The axon in (C) is enlarged for a better illustration. Blue apoptotic signals from cell bodies and axons are indicated with arrowheads. The developmental time is indicated in each image. (D) Z-stack imaging showed the apoptosis of the motor neuron in (C). The axon region in the stack of 13.8 µm is enlarged for a better illustration. The depth of each stack is indicated in each image.

Revised version: Both the cell body and axon of motor neurons can undergo apoptosis during zebrafish early development. (A−D) FRET images of four somites showing motor neurons underwent apoptosis during the development of Tg(mnx1:sensor C3) zebrafish. Two dead motor neurons in (C) are indicated with yellow and white arrowheads respectively. The axon in (D) is enlarged for a better illustration. Blue apoptotic signals from cell bodies and axons are indicated with arrowheads. The developmental time is indicated in each image. (E) Z-stack imaging showed the apoptosis of the motor neuron in (D). The axon region in the stack of 13.8 µm is enlarged for a better illustration. The depth of each stack is indicated in each image.

The apoptosis of motor neurons occurred quickly and almost at the same time between the cell body and axon

Images in Figure 2 showed that blue signals appeared in both the cell body and axon of apoptotic motor neurons. To determine which part of a motor neuron died first and how fast the dying process was, we performed time-lapse imaging of caspase-3 activation in live zebrafish. FRET images in Figure 3A showed that at 32 h 15 min of zebrafish development, both the cell body (indicated with a white arrowhead) and the axon (indicated with a yellow arrowhead) of the indicated motor neuron were green, suggesting that caspase-3 was not activated. Just 3 min later, both the cell body and the axon started to change color (32 h 18 min). Another 3 min later, both the cell body and the axon became blue (32 h 21 min). At 32 h 27 min, two blue apoptotic bodies could be observed in the cell body region and multiple apoptotic bodies were also observed in the axon region. In the next 12 min (until 32 h 39 min), this apoptotic motor neuron gradually degraded into more and smaller apoptotic bodies (Figure 3A; Video S1).

The apoptosis of motor neurons occurred quickly and almost at the same time between the cell body and axon. (A) Real-time imaging of motor neuron apoptosis during zebrafish development. During tracking, one motor neuron changed from green to blue. The cell body is indicated with white arrowheads. The axon is enlarged for a better illustration and indicated with yellow arrowheads. The developmental time is indicated in each image. (B) Real-time imaging of the apoptotic process in another motor neuron during zebrafish development. The cell body is indicated with white arrowheads, and the axon is indicated with yellow arrowheads. The developmental time is indicated in each image.

The time-lapse imaging of another apoptotic motor neuron showed similar results (Figure 3B; Video S2). The color of both the cell body and axon of the motor neuron changed from green to blue within 5 min (from 30 h 45 min to 30 h 50 min). This apoptotic motor neuron gradually degraded into apoptotic bodies in the next 15 min (from 30 h 50 min to 31 h 05 min). These images showed that the activation of caspase-3 can occur within 5-6 min in both the cell body and axon of an apoptotic motor neuron. Thus, we conclude that the cell body and axon of the same motor neuron can undergo apoptosis at the same time during zebrafish development.

Earlier version: Only around 2% of motor neurons died during zebrafish development

Revised version: Only around 2% of motor neurons died during zebrafish early development

Earlier version: In the last several decades, the neurotrophic factor hypothesis has been widely accepted; that is, during the development of vertebrates, approximately 50% of neuronal cells in various regions of the nervous system will die if they cannot receive enough growth factors from their target cells (Castillo-Ruiz, Hite, Yakout, Rosen, & Forger, 2020; Dekkers, Nikoletopoulou, & Barde, 2013; Kandel, Schwartz, Jessell, Siegelbaum, & Hudspeth, 2012; Lossi & Merighi, 2003; Rubin, 1997). Regarding motor neuron development, chick embryos have served as the most studied models. Previous studies showed that half of motor neurons in the spinal cord died between day 5 to day 10 of incubation during the development of chick embryos (Hamburger, 1975; Oppenheim, 1981; Weill, 1991). Similar motor neuron death rates were also obtained between embryonic day 15 to day 18 during the development of rats (Harris & McCaig, 1984; Raoul, Henderson, & Pettmann, 1999; Weill, 1991).

Revised version: In the last several decades, the neurotrophic factor hypothesis has been widely accepted; that is, during the development of vertebrates, approximately 50% of neuronal cells in various regions of the nervous system will die if they cannot receive enough growth factors from their target cells (Castillo-Ruiz, Hite, Yakout, Rosen, & Forger, 2020; Dekkers, Nikoletopoulou, & Barde, 2013; Kandel, Schwartz, Jessell, Siegelbaum, & Hudspeth, 2012; Lossi & Merighi, 2003; Rubin, 1997). Regarding motor neuron development, chick embryos have served as the most studied models. Previous studies showed that half of motor neurons in the spinal cord died between day 5 to day 10 of incubation during the development of chick embryos (Hamburger, 1975; Oppenheim, 1981; Weill, 1991). Similar motor neuron death rates were also obtained in mice (E11.5-E15.5), rats (E15-E18) and humans (11-25 weeks of gestation) during early development (Forger & Breedlove, 1987; Harris & McCaig, 1984; Raoul, Henderson, & Pettmann, 1999; Weill, 1991; Yamamoto & Henderson, 1999). Although the motor neuron death time windows vary in different species in chicks, the common feature of these time windows is that they are all the developmental periods when motor neurons contact with muscle cells. The contact between zebrafish motor neurons and muscle cells occurs before 72 hpf (Kaslin & Ganz, 2020; Myers, Eisen, & Westerfield, 1986; Ott, Diekmann, Stuermer, & Bastmeyer, 2001; Panzer et al., 2005; Schmidt, Strähle, & Scholpp, 2013).

Earlier version: Here, we counted the number of apoptotic motor neurons from 24 hpf to 120 hpf, which covers the major time window of neurogenesis during zebrafish development (Kaslin & Ganz, 2020; Schmidt, Strähle, & Scholpp, 2013) and found that the apoptosis of motor neurons mainly occurred between 24−48 hpf and peaked between 30−36 hpf (Figure 4A). Interestingly, this time window coincided with the time period during which spontaneous muscle twitches occurred in zebrafish (Menelaou et al., 2008; Naganawa & Hirata, 2011). There may be some unknown connection between motor neuron apoptosis and spontaneous muscle twitches, which requires further investigation. Unexpectedly, only a tiny portion (1.6% at 30 hpf and 1.0% at 36 hpf) of the motor neurons underwent apoptosis (Figure 4B), which is much lower than the previous reports that approximately 50% of motor neurons were removed through apoptosis during early development of vertebrates.

Revised version: Thus, we counted the number of apoptotic motor neurons from 24 hpf to 120 hpf, which covers the above time window during zebrafish early development. The data showed that the apoptosis of motor neurons mainly occurred between 24−48 hpf and peaked between 30−36 hpf (Figure 4A). Interestingly, this time window coincided with the time period during which spontaneous muscle twitches occurred in zebrafish (Menelaou et al., 2008; Naganawa & Hirata, 2011). There may be some unknown connection between motor neuron apoptosis and spontaneous muscle twitches, which requires further investigation. Unexpectedly, only a tiny portion (1.6% at 30 hpf and 1.0% at 36 hpf) of the motor neurons underwent apoptosis (Figure 4B), which is much lower than the previous reports that approximately 50% of motor neurons were removed through apoptosis during early development of vertebrates.

Earlier version: Only around 2% of motor neurons died during zebrafish early development. (A) The quantified results show the number of apoptotic motor neurons at each time point during zebrafish early development (n = 30 zebrafish for each time point). (B) The percentages of apoptotic motor neurons among total motor neurons at 30 hpf and 36 hpf. (C) TUNEL assays showing apoptotic motor neurons at 30 hpf and 36 hpf. Immunofluorescence staining with an anti-GFP antibody was used to indicate motor neurons expressing sensor C3 proteins. White arrowheads indicate TUNEL-positive motor neurons. (D) The percentages of apoptotic motor neurons detected by TUNEL assays. (E) The quantified results showing the number of apoptotic motor neurons at 30 hpf after the knockdown of GDNF or BDNF gene using morpholinos (n = 20 zebrafish for each time point). ns, no significant difference.

Revised version: Only around 2% of motor neurons died during zebrafish early development. (A) The quantified results show the number of apoptotic motor neurons at each time point during zebrafish early development (n = 30 zebrafish for each time point). (B) The percentages of apoptotic motor neurons among total motor neurons at 30 hpf and 36 hpf. (C) TUNEL assays showing apoptotic motor neurons at 30 hpf and 36 hpf. Immunofluorescence staining with an anti-GFP antibody was used to indicate motor neurons expressing sensor C3 proteins. White arrowheads indicate TUNEL-positive motor neurons. (D) The percentages of apoptotic motor neurons detected by TUNEL assays.

To validate these results, we performed terminal deoxynucleotidyl transferase dUTP nick end labeling (TUNEL) assays in zebrafish frozen sections. The motor neurons of the sensor zebrafish were identified by immunofluorescence staining using an anti-GFP antibody plus a secondary antibody conjugated with a green dye, and the apoptotic motor neurons were detected by red TUNEL signals (Figure 4C). We analyzed 1,045 motor neurons at 30 hpf, 13 of which were TUNEL-positive, with an apoptotic rate of 1.2%. We also analyzed 652 motor neurons at 36 hpf, 5 of which were TUNEL-positive, with an apoptotic rate of 0.8% (Figure 4D). These results showed that the low apoptotic rates detected by sensor C3-based FRET imaging could be confirmed by the conventional TUNEL assays.

Removed in the revised version: Previous studies showed that neurotrophic factors glial cell-derived neurotrophic factor (GDNF) and brain-derived neurotrophic factor (BDNF) were critical for the survival of motor neurons (Cintrón-Colón, Almeida-Alves, Boynton, & Spitsbergen, 2020; He, Zhang, Yung, Zhu, & Wang, 2013; Kishino, Ishige, Tatsuno, Nakayama, & Noguchi, 1997). Thus, we used morpholinos targeting GDNF and BDNF genes, respectively, to reduce the expression of these two genes. However, we did not observe significant increases in the number of apoptotic motor neurons after the knockdown of these two genes, respectively (Figure 4E).

Most dead motor neurons could not be engulfed by macrophages

Thus far, we have shown that sensor zebrafish can be used to detect motor neuron apoptosis. The subsequent question is how the apoptotic bodies of these dead motor neurons are removed. Macrophages are the major phagocytes to engulf apoptotic bodies in animals (Nagata & Segawa, 2021); therefore, we focused on macrophages. We fully utilized the advantages of sensor zebrafish, which label motor neurons with green fluorescence, and another transgenic zebrafish line, Tg(mpeg1:mCherry), which labels macrophages with a red fluorescent protein (Ellett, Pase, Hayman, Andrianopoulos, & Lieschke, 2011). Live imaging showed that red macrophages distributed everywhere in the zebrafish body at 36 hpf (Figure 5A). Interestingly, after examining 164 apoptotic motor neurons from Tg(mnx1:sensor C3) zebrafish, we only found one case of macrophage engulfment, in which the macrophage indicated with a yellow arrowhead was engulfing the apoptotic bodies of the axon part in the somite, which were indicated with a white arrowhead (Figure 5B). The engulfment percentage was only 0.6% (Figure 5C). These results indicated that the apoptotic bodies of dead motor neurons in the spinal cord were not cleared by macrophages.

Most dead motor neurons could not be engulfed by macrophages. (A) Live imaging showing the distribution of macrophages in zebrafish body at 36 hpf. (B) An example showing the engulfment of apoptotic bodies derived from the axon of a motor neuron by a macrophage in a somite at 36 hpf. The apoptotic bodies are indicated with a white arrowhead and enlarged in the box in the left image. The engulfment by the macrophage is indicated with a yellow arrowhead in the right image. The anatomical structures are indicated. (C) The percentage of apoptotic motor neurons engulfed by macrophages. The total number of apoptotic motor neurons observed, and the number of engulfed apoptotic motor neurons are listed.

Discussion

Earlier version: By expressing FRET-based sensor C3 in zebrafish, for the first time, we were able to monitor the whole process of caspase-3 activation in single motor neurons and the subsequent breakdown of these dead neurons into apoptotic bodies during development. We found that caspase-3 was activated within 5-6 min in an individual apoptotic motor neuron, while the caspase-3 in the rest of the neurons in the same bundle was not activated. More importantly, within the same motor neuron, although the cell body is located in the spinal cord and the axon is extended into the muscle tissue, the caspase-3 activation occurred almost at the same time of observation. These data indicate a synchronized activation pattern of caspase-3 in apoptotic motor neurons.

Revised version: By expressing FRET-based sensor C3 in zebrafish, for the first time, we were able to monitor the whole process of caspase-3 activation in single motor neurons and the subsequent breakdown of these dead neurons into apoptotic bodies during early development. We found that caspase-3 was activated within 5-6 min in an individual apoptotic motor neuron, while the caspase-3 in the rest of the neurons in the same bundle was not activated. More importantly, within the same motor neuron, although the cell body is located in the spinal cord and the axon is extended into the muscle tissue, the caspase-3 activation occurred almost at the same time of observation. These data indicate a synchronized activation pattern of caspase-3 in apoptotic motor neurons.

Added in the revised version: Previous studies showed that some variable primary (VaP) motor neurons died before 36 hpf during zebrafish early development (Eisen, 1992; Eisen, Pike, & Romancier, 1990). Our study examined the overall motor neuron apoptosis rather than a specific type of motor neuron death, so we did not limit our observation to VaP motor neurons. Our data showed sensor zebrafish can successfully capture the death of VaP motor neurons and other types of motor neurons. The reasons are as follows: (1) VaP primary motor neurons die before 36 hpf, but our study found motor neuron cells died after 36 hpf and even at 84 hpf (Figure 4A). (2) The position of the VaP motor neuron is together with that of the caudal primary motor neuron, that is, at the caudal region of the motor neuron cluster. Although it’s rare, we did observe the death of motor neurons in the rostral region of the motor neuron cluster (Figure 2C). (3) There is only one or zero VaP motor neuron in each motor neuron cluster. Although our data showed that usually one motor neuron died in each cluster, we did observe that sometimes more than one motor neuron died in the motor neuron cluster (Figure 2C).

Earlier version: It is generally accepted that during the development of vertebrates, approximately half of the motor neurons are removed by apoptosis. However, our results showed that motor neurons were not prone to death during early zebrafish development: only around 2% of motor neurons died of apoptosis during the early development of zebrafish. This large discrepancy may be the unique properties of zebrafish compared with chicks and rats, which could be further investigated. Another possibility is that the limitations of conventional apoptosis detection methods such as the terminal deoxynucleotidyl transferase dUTP nick end labeling (TUNEL) assay (Loo, 2011), nuclear morphology staining (H. Li et al., 2000) and caspase-3 staining (Wang et al., 2010) may potentially exaggerate the number of apoptotic cells in tissue sections. For example, the specificity of antibodies for apoptotic markers or the quality of other reagents in experiments may not be good enough, which can result in false positive apoptotic signals. Besides, these staining methods have difficulties in distinguishing neuronal cells and nearby glial cells; thus, dead glial cells may also be counted as dead neuronal cells. GDNF and BDNF are well-known neurotrophic factors for motor neurons. However, our results showed that inhibition of these two genes could not increase the death rates, which also indicated that zebrafish motor neurons were not prone to cell death during early development.

Revised version: It is generally accepted that during the development of vertebrates, approximately half of the motor neurons are removed by apoptosis. However, our results showed that motor neurons were not prone to death during zebrafish early development: only around 2% of motor neurons died of apoptosis during the early development of zebrafish. This large discrepancy may be the unique properties of zebrafish compared with chicks and rats, which could be further investigated. Another possibility is that the limitations of conventional apoptosis detection methods such as the terminal deoxynucleotidyl transferase dUTP nick end labeling (TUNEL) assay (Loo, 2011), nuclear morphology staining (H. Li et al., 2000) and caspase-3 staining (Wang et al., 2010) may potentially exaggerate the number of apoptotic cells in tissue sections. For example, the specificity of antibodies for apoptotic markers or the quality of other reagents in experiments may not be good enough, which can result in false positive apoptotic signals. Besides, these staining methods have difficulties in distinguishing neuronal cells and nearby glial cells; thus, dead glial cells may also be counted as dead neuronal cells.

Earlier version: The removal of dead cells by phagocytes can effectively reduce the occurrence of inflammation, which is very important for animals (Arandjelovic & Ravichandran, 2015). Our study showed that macrophages, although serving as professional phagocytes, could not enter the spinal cord to engulf the apoptotic motor neurons. Thus, some other cells in the spinal cord may help to engulf the dead cells. For example, a recent zebrafish study reported that neural crest cells can help to remove apoptotic cells in the nervous system during development (Zhu et al., 2019). In the future, we will investigate whether red fluorescent protein-labeled neural crest cells could engulf dead motor neurons in our sensor zebrafish.

Revised version: The removal of dead cells by phagocytes can effectively reduce the occurrence of inflammation, which is very important for animals (Arandjelovic & Ravichandran, 2015). Our study showed that macrophages, although serving as professional phagocytes, did not enter the spinal cord to engulf the apoptotic motor neurons. Thus, some other cells in the spinal cord may help to engulf the dead cells. For example, a recent zebrafish study reported that neural crest cells can help to remove apoptotic cells in the nervous system during development (Zhu et al., 2019). In the future, we will investigate whether red fluorescent protein-labeled neural crest cells could engulf dead motor neurons in our sensor zebrafish.

During the past several decades, the real-time tracking of neuronal cell apoptosis under physiological conditions in live animals remains challenging. Since the development of fluorescent proteins, several FRET-based protein biosensors have been designed for the detection of apoptosis in live cells cultured in petri dishes (Ai, Hazelwood, Davidson, & Campbell, 2008; Albeck et al., 2008; Karasawa, Araki, Nagai, Mizuno, & Miyawaki, 2004; Mahajan, Harrison-Shostak, Michaux, & Herman, 1999; Rehm et al., 2002; Tyas, Brophy, Pope, Rivett, & Tavaré, 2000; Xu et al., 1998). However, unlike the good performances of these biosensors in vitro, the in vivo applications of these biosensors have been unsatisfactory because of their small dynamic ranges and the complex optical environments in live animals (Andrews et al., 2016; Takemoto et al., 2007; Takemoto, Nagai, Miyawaki, & Miura, 2003; Yamaguchi et al., 2011). We compared sensor C3 with other FRET-based apoptotic biosensors and found that sensor C3 can achieve a FRET effect of 4- to 5-fold, while other biosensors usually had FRET effects of less than or near 2-fold. This high FRET effect makes sensor C3 the ideal biosensor for visualizing neuronal cell apoptosis in live animals where other biosensors encounter difficulties.

In summary, our results showed that sensor zebrafish can serve as valuable tools for studying motor neuron apoptosis in live zebrafish, which could improve our understanding of neuronal cell apoptosis during normal development and in neurodegenerative diseases.

Materials and methods

Zebrafish maintenance

Zebrafish (AB strain) were raised at 28 °C in a ZebTEC system. The photoperiod was 14 h of light and 10 h of darkness. Zebrafish were fed with artemia twice a day. The experiments in this study were approved by the Animal Research Ethics Committee of the University of Macau (Protocol ID: UMARE-032-2016).

Generation of sensor zebrafish

The zebrafish genome was extracted using the DNeasy Blood & Tissue Kit (Qiagen, 69504). A zebrafish motor neuron-specific promoter mnx1 was amplified by polymerase chain reaction (forward primer, ACGCGTCGACGAATTCATTTAAATTAGCCTGGCATC, reverse primer, CCCACCGGTCTGGCCCACCTCACAAACAGATTA). The promoter sequence and sensor C3 gene were cloned into the same plasmid backbone to generate the transgenic vector, named pSK-mnx1-C3. Tol2 mRNA was obtained by in vitro transcription using the mMESSAGE mMACHINE SP6 transcription kit (Ambion, AM1340). The transgenic vector containing the sensor C3 gene (final concentration, 50 ng/μL) and Tol2 mRNA (final concentration, 100 ng/μL) were mixed and injected (2-4 nL) into newly fertilized zebrafish eggs using an MPPI-3 pressure injector. The zebrafish embryos with green fluorescence indicating sensor C3 expression were selected under a stereomicroscope with a GFP filter. These positive zebrafish were cultured to sexual maturity and crossed with wild-type zebrafish to obtain transgenic sensor zebrafish.

In vivo FRET imaging of neuronal cell apoptosis

Sensor zebrafish embryos or larvae were anesthetized with 0.016% tricaine (Sigma-Aldrich, E10521) solution. Anesthetized zebrafish were embedded in low melting point agarose gel (Promega, V2111) in a confocal dish. The zebrafish were then imaged using a Carl Zeiss LSM 710 or LSM880 confocal laser scanning microscope. For FRET imaging, a 458 nm laser was used to excite sensor C3 proteins, and the emissions were collected simultaneously in two channels: 460-500 nm was collected as the CFP channel, designated as blue, and 520-550 nm was collected as the YFP channel, designated as green. Then, the CFP and YFP images were merged to generate the FRET images. In live neuronal cells, sensor C3 emitted green fluorescence when CFP was excited because of the energy transfer from CFP to YFP, while in apoptotic neuronal cells, activated caspase-3 proteins cleaved sensor C3 at the DEVD site, and the FRET effect was abolished. Thus, sensor C3 emitted blue fluorescence when CFP was excited under the same condition.

Immunofluorescence and TUNEL assay

Earlier version: Tg(mnx1:sensor C3) zebrafish were fixed in 4% paraformaldehyde at 4 °C for 12 h. The fixed zebrafish were dehydrated in 30% sucrose at 4 °C for 48 h. The zebrafish were then embedded in Shandon Cryomatrix embedding resin (Thermo Scientific, 6769006), and frozen sections were cut at 25 µm using a Leica CM3050S cryostat. The sections were washed with phosphate-buffered saline (PBS) three times to remove the embedding resin and then were blocked with 1% bovine serum albumin with 0.1% Triton X-100 in PBS for 1 h at room temperature (RT). Then, the sections were coated with anti-GFP primary antibody (Cell Signaling, 2956, 1:400) for 2 h at RT. After that, the sections were washed with PBS three times and then coated with Alexa Fluor 488-conjugated goat anti-rabbit secondary antibody (Thermo Scientific, A-11034, 1:200) for 1 h at RT. The sections were washed with PBS three times. The terminal deoxynucleotidyl transferase and Cy3-dUTP labeling solution in the TUNEL kit (Beyotime, C1089) were mixed. The sections were then incubated with this reaction mixture for 1 h at 37 °C. After the incubation, the sections were washed with PBS three times and then mounted with Mowiol mounting medium (Millipore, 475904). The sections were imaged with a Carl Zeiss LSM880 confocal laser scanning microscope with 488 nm and 561 nm lasers.

Revised version: Tg(mnx1:sensor C3) zebrafish were fixed in 4% paraformaldehyde at 4 °C for 12 h. The fixed zebrafish were dehydrated in 30% sucrose at 4 °C for 48 h. The zebrafish were then embedded in Shandon Cryomatrix embedding resin (Thermo Scientific, 6769006), and frozen sections were cut at 25 µm using a Leica CM3050S cryostat. The sections were washed with phosphate-buffered saline (PBS) three times to remove the embedding resin and then were blocked with 1% bovine serum albumin with 0.1% Triton X-100 in PBS for 1 h at room temperature (RT). Then, the sections were coated with anti-GFP primary antibody (Cell Signaling, 2956, 1:400) for 2 h at RT. After that, the sections were washed with PBS three times and then coated with Alexa Fluor 488-conjugated goat anti-rabbit secondary antibody (Thermo Scientific, A-11034, 1:200) for 1 h at RT. The sections were washed with PBS three times. The terminal deoxynucleotidyl transferase and Cy3-dUTP labeling solution in the TUNEL kit (Beyotime, C1089) were mixed.

Added in the revised version: The sections were then incubated with this reaction mixture for 1 h at 37 °C. After the incubation, the sections were washed with PBS three times and then mounted with Mowiol mounting medium (Millipore, 475904). The sections were imaged with a Carl Zeiss LSM880 confocal laser scanning microscope with 488 nm and 561 nm lasers.

Removed in the revised version: Microinjection of morpholinos

Removed in the revised version: Zebrafish eggs were collected after fertilization and were placed on an agarose plate. Approximately 5 nL morpholino (0.5 mM) specific to the GDNF gene (sequence, TGGCTAGAATGTCCCATAACTTCAT) or BDNF gene (sequence, TGGTCATCACTCTTCTAACCTGTTG) of zebrafish was injected into one-cell stage embryos using the MPPI-3 pressure injector. After injection, the zebrafish were quickly transferred into fresh fish water.

In vivo imaging of macrophage engulfment

Tg(mnx1:sensor C3) zebrafish were crossed with Tg(mpeg1:mCherry) zebrafish to get offspring that neuronal cells and macrophages were labeled with green and red, respectively. The zebrafish were then imaged using a Carl Zeiss LSM880 confocal laser scanning microscope. FRET imaging was applied to detect apoptotic neuronal cells. At the same time, a 561 nm laser was used to excite the mCherry proteins for the visualization of macrophages.

Statistical analysis

All of the data are expressed as the means ± SD. Statistical significance was judged by one- or two-way ANOVA using GraphPad Prism 7 software.